```python
NAME = "Bryan Tchakote"
```

---

```python
# Lines that start with # are comments.  They won't "run" as code.  We can use them to explain things.
# This cell will produce an error because of the >>> in front of the code! Try and see when you run the cell.
>>> print("test")
```

Output:
```
File "C:\Users\tchak\AppData\Local\Temp\ipykernel_132716\1398651531.py", line 3
    >>> print("test")
    ^
SyntaxError: invalid syntax
```

This is correct:

```python
print("test")
```

Output:
```
test
```

An example of creating and using a function in the notebook, as you will do for the homework, is here:

```python
# Here I'm defining a function that takes and argument and prints hello. 
# When I run this cell, nothing visible happens, but the function is stored in memory so it will run when I call it.
# If you don't run this cell, you will get an error on the next cell because it was not defined!
def print_hi(name):
    print("Hi,", name)
```

```python
# Now I am calling it with my name as the "name" argument.  It should print the greeting when you run this cell!
print_hi("Alice")
```

Output:
```
Hi, Alice
```

**Here are some key points on how to use the notebook and submit your work.**
We will grade based on assuming you have read and understood them.
    
1. **Using the Notebook to Show Your Work**: You must learn to write code in the notebook... It is a core tool for data science and will make it easier to develop and document your work if you become good at using it. Writing your code in another tool and pasting it into the notebook will probably not work well (forgetting to include code elsewhere, messing up the spacing, or code that don't run as copied). You must be sure your code cells execute because we will test them.  So learn how to run code in the notebook cells to double check your work.
2. **Read Directions Carefully**:  The instructions for your code are very important. If you don't follow the requirements, your application won't do as requested. Making it work correctly is part of learning to program.  If we worded something unclearly, ask the teacher.

The second set of issues are coding style:

1. **Indentation**: In Python, indentation matters and must be consistent. If you write your code in the Notebook, the **tab** key will indent properly.  If you use another editor and paste into the notebook, it might not be correctly indented (when you do write code in another editor, make sure you set your tabs to indent as 4 spaces, not as a tab character.)  You must make sure that your pasted code runs in the notebook or it will not get a good grade. Anyway, we recommend beginners to work in a Jupyter notebook for this course, whether it's this one or a draft file.
2. **Spacing**: Follow closely the spacing shown in the lessons. There should not be a space  between a function name and the parentheses with the arguments. As a programmer, style is very important. If you work with programmers in the future, they sometimes have "lint" checkers to test your code for style and reject if it doesn't follow the approprate spacing and blank-line-rules. Think of it as a matter of politness for other people reading your code \ (•◡•) /
3. **Names of Variables**: In Python, there's a culture of making everything readable. Don't use ``x`` and ``y`` as your variable names... use words like ``pounds`` and ``kilograms``. It will be easier for colleagues (and yourself) to understand the code later.
4. **Error Messages**: Please use informative error messages that tell the user what they did wrong and what kind of input you expect. Imagine you are designing the user experience! Think about how to help your user. And remember **you** are the user when you debug!

We will take points off for issues of non-standard spacing, indentation, bad error messages, and bad variable names in the future.  This will continue for the entire course.

There are multiple ways to code all the answers.  Here are a few more code style tips:

1. If you do a calculation or a transformation, like ``float(pounds)`` -- do it once and save it as a variable, don't do it multiple times.  You should try not to have code that repeats itself too much.  If you repeat things, you can make mistakes like typos and it will be harder to find them. Also, it's wasting computer power.
2. Tests like "4 < test < 40" need to be saved in a variable or used in a ``if`` statement.  It won't do anything relevant otherwise.
3. ``try``/``except`` should be used to catch errors. (In fancier, more formal Python, there is more careful error catching where the type of error is detected and handled. We're just doing the basic try/except right now.) Anytime you have a conversion or something that could result in an error, you should wrap it in try/except. Do not allow a user to run code that results in an un-handled error.

### Question 0 ###

Follow the directions in each cell's comment for what to do here.

<!-- grade_id: var_float, points: 0, grade: false, solution: true -->
```python
# Create a variable named "myvar" and assign it the value of the float version of 6 (hint, decimals):

myvar = 6.
```

<!-- grade_id: test_var_float, points: 1, grade: true, solution: false -->
```python
# Don't change or try to delete this cell, it's a hidden test for your code.

```

<!-- grade_id: var_string, points: 0, grade: false, solution: true -->
```python
# Now create a variable named "my_greeting" and assign it a string value for Bonjour (hint: how do you 
# assign a string to a variable?):

my_greeting = "Bonjour"
```

<!-- grade_id: test_var_string, points: 1, grade: true, solution: false -->
```python
# Don't change this cell. It's a hidden test for your answer.

```

### Question 1:

Complete the short function below which should prompt the user for a weight in kilograms (using the ``input`` function) and convert it to a weight in UK stones. (You will have to Google for the conversion formula.)  It should print out the weight in UK stones in a nice message. It should also return the value after printing. Don't round (or truncate) the result, or our tests won't pass and you might not get credit.

Show it working with your input. Create a new cell or cells using the + symbol in the notebook to show it working.  Show it taking your input and returning the answer.  There are some test cases here for your code using "assert" that must work properly.

<!-- grade_id: convert_simple, points: 0, grade: false, solution: true -->
```python
def convert():
    ''' Ask for the user input with a prompt and compute the weight in kilograms. Print the value
    and return it from the function as well.'''
    
    prompt = "Enter a weight in kilograms: "
    stones = float(input(prompt)) * .157473
    print(f"Your input is equal to {stones} stones")
    return stones
```

To run your code in a cell, you should add a new cell and type:

```convert()```

and execute that cell.  Provide the input and show it working below:

```python
# run this to show it working!
convert()
```

Output:
```
Enter a weight in kilograms: 34
Your input is equal to 5.354082 stones
5.354082
```

The following code checks that your function performs correctly.  It is only verifying the return value is the correct number.  Testing functions for their return values is a standard practice in *Test-Driven Development*.  Test cases help you figure out where your code has errors, if you have lots of functions calling each other.

<!-- grade_id: test_convert_simple, points: 2, grade: true, solution: false -->
```python
# This is a test of your code. Don't change this.

import mock
from io import StringIO
import sys

def test_convert():
    with mock.patch('builtins.input', return_value='34'):
        assert convert() == 5.354082
    with mock.patch('builtins.input', return_value='54'):
        assert convert() == 8.503542
            
test_convert()
```

Output:
```
Your input is equal to 5.354082 stones
Your input is equal to 8.503542 stones
```

### Question 2: ###

Update the code from question 1 to handle errors in input. Copy the code and modify it in the new function below.

* to include try / except to handle non-numeric input. If the user inputs non-numeric input (for instance, a string) or negative value or a 0, print an error message that says exactly: ``"Error: You must enter a positive number!"`` and don't return anything.
* return the resulting number of stones from the function, if it does not produce an error. If you want to round the output to 2 decimal places, that would be extra.

Hints: You can put multiple statements after your ``try:`` part, before the ``except:``.  

<!-- grade_id: better_convert, points: 0, grade: false, solution: true -->
```python
def better_convert():
    """ Your function should prompt for the input. Also check for
    valid numeric positive input (not a string and not 0) and give error feedback if it's not 
    valid. If it's not valid, return without printing the answer and don't return an answer.
    Print the value in stones if the input was a valid number and return the result.
    """
    
    err_msg = "Error: You must enter a positive number!"
    
    prompt = "Enter a weight in kilograms: "
    try:
        stones = float(input(prompt)) * 0.157473
        assert stones > 0
    except (ValueError, AssertionError):
        print(err_msg)
        return
    
    print(f"Your input is equal to {round(stones, 2)} stones")
    return stones
```

```python
# run this to show it is working!
better_convert()
```

Output:
```
Enter a weight in kilograms: 34
Your input is equal to 5.35 stones
5.354082
```

<!-- grade_id: test_better_convert1, points: 1, grade: true, solution: false -->
```python
# testing input...

with mock.patch('builtins.input', return_value='34'):
    with mock.patch('sys.stdout', new_callable=StringIO):
        assert better_convert() == 5.354082
        assert '5.35' in sys.stdout.getvalue()
```

<!-- grade_id: test_better_convert2, points: 1, grade: true, solution: false -->
```python
# testing more....

with mock.patch('builtins.input', return_value='cat'):
    with mock.patch('sys.stdout', new_callable=StringIO):
        assert better_convert() is None
        assert sys.stdout.getvalue() == "Error: You must enter a positive number!\n"
```

<!-- grade_id: test_better_convert3, points: 1, grade: true, solution: false -->
```python
# testing even more....

with mock.patch('builtins.input', return_value='-3'):
    with mock.patch('sys.stdout', new_callable=StringIO):
        assert better_convert() is None
        assert sys.stdout.getvalue() == "Error: You must enter a positive number!\n"
```

<!-- grade_id: test_better_convert4, points: 1, grade: true, solution: false -->
```python
# testing even more....

with mock.patch('builtins.input', return_value='0'):
    with mock.patch('sys.stdout', new_callable=StringIO):
        assert better_convert() is None
        assert sys.stdout.getvalue() == "Error: You must enter a positive number!\n"
```

### Question 3: ###

Write 2 functions.  The second function will use the output from the first function.

The first function take as arguments 2 integers, and returns a random number between those integers. We will call this function ``get_rand_int``. If the two integers are equals, this is an error and we will return None (hint: just put ```return None``` as your line for these cases). We can test for the result of None in the second function and report that it was an error there.

The second function, called ``print_random``, asks for user input for one integer greater than 0, and if it receives good input, calls the first function.  It converts the user input to an integer and calls the first function with 0 as the first argument and the user input as the second argument.  We capture the output from the first function in a variable (e.g., ``rand = get_rand_int(0, input))``.  Our second function then prints out the number returned by the first function in a sentence: "Your random number is " and then the number.

It should do error checking and print:
"Invalid input: You must input an integer greater than 0 for your input." if the user input is not convertable to a valid integer greater than 0.  You should also check if the result from the ``get_rand_int`` function is ``None`` or not.

Finally, return the value of the random number as the value of the second function.

**Note**

The function ``randint`` from the ``random`` module may be useful here. It was introduced in the book!


```python

import random

help(random.randint)
```

Output:
```
Help on method randint in module random:

randint(a, b) method of random.Random instance
    Return random integer in range [a, b], including both end points.
```

<!-- grade_id: two_functions, points: 0, grade: false, solution: true -->
```python
def get_rand_int(arg1, arg2):
    """ Takes 2 integer arguments, and returns a random integer between them.
    Return None if the arguments aren't 2 different integers.
    """
    
    if isinstance(arg1, int) and type(arg2) == int and arg1 != arg2:
        return random.randint(arg1, arg2)
    
def print_random():
    """ Prompt the user to input an integer > 0, and return a random integer between 0 and
    their input value.  Print the required error if their input is not appropriate.
    """
    
    prompt = "Please enter an integer greater than 0: "
    err_msg = "Invalid input: You must input an integer greater than 0 for your input."
    
    try:
        number = int(input((prompt)))
        assert number > 0
    except (ValueError, AssertionError):
        print(err_msg)
        return
    
    rand_num = get_rand_int(0, number)
    print(f"Your random number is {rand_num}")
    return rand_num
```

```python
# show this running... 
print_random()
```

Output:
```
Please enter an integer greater than 0: 5
Your random number is 3
3
```

<!-- grade_id: test_get_rand_int1, points: 4, grade: true, solution: false -->
```python
# This is a test of your code. First we test the get_rand_int function.

def test_get_rand_int():
    fail_count = 0
    if get_rand_int(0, 0) is None:
        fail_count += 1
    if get_rand_int('car', 0) is None:
        fail_count += 1
    if get_rand_int(1.5, 5) is None:
        fail_count += 1
    if get_rand_int(0, 1.5) is None:
        fail_count += 1
    if get_rand_int(5, 5) is None:
        fail_count += 1
    print(fail_count)  # possible partial points?
    return fail_count

assert test_get_rand_int() == 5
```

Output:
```
5
```

<!-- grade_id: test_print_rand1, points: 1, grade: true, solution: false -->
```python
# This is a test of your second function.

with mock.patch('builtins.input', return_value='-1'):
    with mock.patch('sys.stdout', new_callable=StringIO):
        result = print_random()
        assert not result
        assert sys.stdout.getvalue() == "Invalid input: You must input an integer greater than 0 for your input.\n"
```

<!-- grade_id: test_print_rand2, points: 1, grade: true, solution: false -->
```python
# This is a test of your second function.

with mock.patch('builtins.input', return_value='0'):
    with mock.patch('sys.stdout', new_callable=StringIO):
        result = print_random()
        assert not result
        assert sys.stdout.getvalue() == "Invalid input: You must input an integer greater than 0 for your input.\n"
```

<!-- grade_id: test_print_rand3, points: 1, grade: true, solution: false -->
```python
# This is a test of your second function.

with mock.patch('builtins.input', return_value='5'):
    with mock.patch('sys.stdout', new_callable=StringIO):
        result = print_random()
        assert sys.stdout.getvalue() != "Invalid input: You must input an integer greater than 0 for your input.\n"
        assert result <= 5 and result >= 0
```

<!-- grade_id: test_print_rand4, points: 1, grade: true, solution: false -->
```python
# This is a test of your second function.

with mock.patch('builtins.input', return_value='cat'):
    with mock.patch('sys.stdout', new_callable=StringIO):
        result = print_random()
        assert sys.stdout.getvalue() == "Invalid input: You must input an integer greater than 0 for your input.\n"
```

### Question 4: ###

If I have an original placement $P$ in an account with a nominal annual interest rate $r$, after $n$ years, the total amount of **extra money** I get (the compound interest) is given by:

$$I = P \times ((1 + r) ^ n - 1)$$

Write a function ``compound_interest`` that takes the original placement, the nominal annual interest rate and a number of years as parameters and *returns* (not print) the compound interest after these number of years.

Use the ``assert`` statement and the ``isinstance`` function to check that $n$ is a strictly positive integer, $P$ is positive and $r$ is positive.

Make sure there's a docstring attached to the function that specifies that

- $P$ is the original placement
- $r$ is the nominal interest rate
- $n$ is the number of years the compound interest is computed on

A docstring is a simply a string right after the function definition that documents the function. It's displayed by the ``help()`` builtin:

```python
# Example showing a docstring -- the help you create for a function.

def smile():
    'prints a large smiley on screen'  # <- this is a docstring explaining what the function does.
    print('^___________^')
    
help(smile)  # <- this displays the help associated to the smile function
```

Output:
```
Help on function smile in module __main__:

smile()
    prints a large smiley on screen
```

```python
# Example of a longer doc string. The convention is 3 quote marks.
def smile2():
    ''' This is a really long docstring, that takes multiple rows.
    This one prints out a smaller ascii smiley.
    '''
    print(":)")
    
help(smile2)
```

Output:
```
Help on function smile2 in module __main__:

smile2()
    This is a really long docstring, that takes multiple rows.
    This one prints out a smaller ascii smiley.
```

<!-- grade_id: cell-b1c40b1491cb0e1e, points: 0, grade: false, solution: true -->
```python
def compound_interest(placement, rate, years):
    """
    Returns the compound interest of an original placement given the annual interest rate and the number of years
    
    Args:
        placement (float): The original placement.
        rate (float): The nominal interest rate.
        years (int): The number of years the compound interest is computed on.
        
    Returns:
        float: The compound interest produced by the placement.
    """

    assert placement >= 0 and rate >= 0 and isinstance(years, int) and years > 0    
    return placement * ((1 + rate)**years - 1)
```

<!-- grade_id: test_comp_inputs, points: 2, grade: true, solution: false -->
```python
# testing input and return value.
# The first test makes looks strange to you, we can talk about accuracy issues during
# the face-to-face session if you're brave.  This is a hard thing to write a good test for.

assert abs(compound_interest(100, .01, 1) -1) < 10e-5

```

<!-- grade_id: test_comp_res2, points: 3, grade: true, solution: false -->
```python
# testing assert-based input preconditions

# These tests use fancier try/except, which we can discuss in class.

```

<!-- grade_id: cell-d3c24f97b6e1ed38, points: 1, grade: true, solution: false -->
```python
# testing presence of a non-empty docstring that contains the words **placement** and **rate**
help(compound_interest)


```

Output:
```
Help on function compound_interest in module __main__:

compound_interest(placement, rate, years)
    Returns the compound interest of an original placement given the annual interest rate and the number of years
    
    Args:
        placement (float): The original placement.
        rate (float): The nominal interest rate.
        years (int): The number of years the compound interest is computed on.
        
    Returns:
        float: The compound interest produced by the placement.
```